```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

<table style="width: 100%;">
    <tr style="background-color: transparent;"><td>
        <img src="https://data-88e.github.io/assets/images/blue_text.png" width="250px" style="margin-left: 0;" />
    </td><td>
        <p style="text-align: right; font-size: 10pt;"><strong>Economic Models</strong>, Fall 2025<br>
            Dr. Eric Van Dusen <br>
            Notebook by: Umar Maniku <br>
            Rohan Jha <br>
            Akhil Venkatesh</p></td></tr>
</table>

# Lab 4: The Cobb-Douglas Production Function

```python
from utils import *
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd 
import datascience
import seaborn as sns
import plotly.graph_objs as go
from plotly.offline import download_plotlyjs, init_notebook_mode, plot, iplot
from ipywidgets import interact, interactive, fixed, interact_manual
import ipywidgets as widgets
from IPython.display import display
#import nbinteract as nbi
import warnings
warnings.filterwarnings('ignore')
%matplotlib inline
plt.style.use("seaborn-v0_8-muted")
```

## Part 1: Production and Cobb-Douglas Functions in Macroeconomics

### The Cobb-Douglas Production Function

The Cobb-Douglas Production Function is as follows:
$$
f(K, L) = K^\alpha L^\beta \\
Y = A \cdot f(K, L) = A K^\alpha L^\beta
$$
where $\alpha$ and $\beta$ are exponents.

A common simplification is that $\beta = 1 - \alpha$. We will later explore the implications of this statement. For now, let us rewrite the above function:

$$
Y = A K^\alpha L^{1 - \alpha}
$$

Note that this is a function of two variables, $K$ and $L$. If we were to plot this function utilizing both variables, we would need a 3D plot with $K$, $L$ and $Y$ each having their own axis. For now, let us gain greater insight into what this function would look like by holding one variable constant and plot the other versus output.

For the first case, let us visualize the Cobb-Douglas Production Function with output as a function of capital, holding the amount of labor constant at $\bar L$.

**Question 1.1** Define ```cobb_douglas```, a function that calculates output using the Cobb-Douglas function. This should work for constant, decreasing and increasing returns to scale. We have provided an ```if-else``` structure to faciliate your implementation. Note that the syntax `beta = None` tells Python that the default argument for the `beta` argument is the value `None`. This is not a number, so you'll need to handle this case.

_Hint:_ When do we substitute $\beta = 1 - \alpha$?


```python
def cobb_douglas(A, K, L, alpha, beta = None):
    if alpha is not None:
        return A * K ** alpha * L ** (1 - alpha)
    else:
        return A * K ** alpha * L 
```

```python
grader.check("q1_1")
```

Output:
```
q1_1 results:
    q1_1 - 1 result:
        ✅ Test case passed

    q1_1 - 2 result:
        ❌ Test case failed
        Trying:
            np.isclose(cobb_douglas(1, 4, 0.5, 0.2, 1.1), 0.6155722066724582)
        Expecting:
            True
        **********************************************************************
        Line 1, in q1_1 1
        Failed example:
            np.isclose(cobb_douglas(1, 4, 0.5, 0.2, 1.1), 0.6155722066724582)
        Expected:
            True
        Got:
            False

    q1_1 - 3 result:
        ✅ Test case passed
```

<!-- BEGIN QUESTION -->

**Question 1.2** Now let's plot Cobb-Douglas as a function of the capital stock, holding labor constant at $\bar L$. To do this, we'll need to write a function called `cobb_douglas_plotter_K` that takes in the Cobb-Douglas model parameters to create an array of values for $Y$ called `Y_s`. In the function body, fill in the correct values for `K_s` and `Y_s`. `K_s` should be an array of the possible values of $K$ in the range $[0, 1]$ using intervals of 0.01. When you have completed the function, call it using $A = 1$, $\bar L = 0.5$, and $\alpha = 0.4$.

_Hint:_ Your `cobb_douglas` function can take any of its parameters as an array.


```python
def cobb_douglas_plotter_K(A, L_bar, alpha):
    K_s = ...
    Y_s = ...
    
    ### Do not edit the code below; this does the plotting ###
    plt.plot(K_s, Y_s)
    plt.title(fr"Cobb-Douglas with $\bar L$ = {L_bar}, $A$ = {A} and $\alpha$ = {alpha}")
    plt.xlabel("Capital Stock")
    plt.ylabel("Output");


...
```

<!-- END QUESTION -->

If you defined the function correctly, you should notice some of the properties in the graph above:
1. It is increasing. This is called increasing returns to capital wherein any increase in capital will lead to an increase in output, assuming that labor is held constant.
2. It is concave (increasing at a decreasing rate). This is called diminishing marginal returns to capital wherein any additional unit of capital will lead to smaller and smaller increases in capital.

We can do the same thing but instead holding capital constant and letting labor vary.

```python
cobb_douglas_plotter_L(1, 0.5, 0.4)
```

We notice that the same properties hold: the function is exhibiting both increasing returns to labor and diminishing marginal returns to labor. This is because we selected the same parameters as when we plotted output versus capital stock. The cell below has the same function call, but try using different parameters and observe what happens to its shape.

```python
# Change the parameters below
# First parameter: alpha, Second parameter: Fixed value of K, Third parameter: beta

cobb_douglas_plotter_L(2, 0.9, 0.7)
```

<!-- BEGIN QUESTION -->

**Question 1.3** Suppose a country doubles its capital stock while holding labor constant. Using the Cobb-Douglas production function $$Y = AK^{\alpha}L^{1-\alpha}$$ explain why output increases but at a diminishing rate. Support your answer by discussing the behavior of the marginal product of capital (MPK).

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Implications for cross-country comparisons

Work by Professors C.W. Cobb and P.H. Douglas found that production or output was a weighted average of the log of capital and labor. The equation for Cobb-Douglas production functions were the result of their research, especially when a log transformation was applied to the equation:

$$\begin{aligned}
Y &= A K^\alpha L^{1 - \alpha} \\
\ln Y &= \ln A + \alpha \ln K + (1 - \alpha) \ln L
\end{aligned}
$$

Note that this is exactly the weighted average that the two Professors found in their empirical findings: capital and labor are weighted by $\alpha$ and $1 - \alpha$ respectively. However, this is still showing production as a function of two variables, $K$ and $L$. Rearranging the equation yields something interesting:

$$\begin{aligned}
\ln Y &= \ln A + \alpha \ln K + \ln L - \alpha \ln L \\
\ln Y- \ln L &= \ln A + \alpha \ln K - \ln L \\
\ln \dfrac{Y}{L} &= \ln A + \alpha \ln \dfrac{K}{L}
\end{aligned}
$$

The Cobb-Douglas function is now an equation in 1 variable: $\ln \frac{K}{L}$, making it linear instead of exponential.

Now that you've identified a process of capturing all parts of the Cobb-Douglas function, we are now able to compare  values of $A$ and $\alpha$ between countries, and by extension how capital and labor are deployed in different ways between nations. We will now learn about the significance of $A$ and $\alpha$ in the sections below. The upcoming project will have you try this on a broad range of nations from 1990 to 2019.

### Shifts in $A$ and its effect on output

```python
A_slider = widgets.FloatSlider(min = 0.5, max = 10, step = 0.5, value = 1)
A_shifter = widgets.interactive(change_A, A = A_slider)
display(A_shifter)
```

```python
A_shifter.close()
```

Supply or total factor productivity shocks could cause $A$ to change. These occur if there is a change in total output for a given level of capital and labor. Examples of these include financial crises, technology shocks, natural environment/disasters and energy prices. 

Favorable shocks rotate the production function upward through an increase in $A$. Thus, each unit of input from capital and labor now simultaneously produce more output. What does this mean for the rental rate of capital and the real wage? Recall the functions for both of them:

$$\begin{aligned}
MPL &= \dfrac{\partial Y}{\partial L} \\
&= A (1 - \alpha) \left ( \dfrac{K}{L} \right )^{\alpha} \\
MPK &= \dfrac{\partial Y}{\partial K} \\
&= \alpha A \left ( \dfrac{L}{K} \right )^{1 - \alpha} \\
\end{aligned}
$$

Both MPK and MPL will increase by a factor of $A$. Thus, it would be more expensive to hire an additional unit of labor or rent an additional unit of capital. As they are both more productive than they previously were, they are both more valuable to a business and thus will cost more.


Negative shocks do the opposite. They rotate the production function downward through a decrease in $A$. Each unit of input is now less productive, meaning that both the rental rate of capital and the real wage are lower.

<!-- BEGIN QUESTION -->

**Question 1.4.** Suppose a technological innovation doubles total factor productivity ($A$). Explain how this shock affects both (i) the output level for a given $K$ and $L$, and (ii) the marginal products of capital and labor. Use the Cobb-Douglas function and textbook reasoning in your explanation.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Differences in $\alpha$ and their effects on output

```python
alpha_slider = widgets.FloatSlider(min = 0.1, max = 1, step = 0.1, value = 0.5)
alpha_wid = widgets.interactive(change_alpha, alpha = alpha_slider)
display(alpha_wid)
```

```python
alpha_wid.close()
```

$\alpha$ and $\beta$ are called the output elasticities of capital and labor, respectively. They measure the responsiveness of output to a change in the levels of either labor or capital, holding all else constant. This means that if $\alpha$ or $\beta$ were high, then any small increase in their respective input would lead to a relatively large increase in output. As an example, if $\alpha$ were 0.4, then a 1% increase in capital would lead to a 0.4% increase in output.

**Question 1.5.** Let's link this back to Cobb-Douglas' simplifying assumption of constant returns to scale. Recall that this means $\alpha = 1 - \beta$. If a nation has a high $\alpha$, what does this mean about how resources are allocated?

<ol type="A" style="list-style-type: lower-alpha;">
    <li>The higher the $\alpha$, the more efficient a country's production is.</li>
    <li>The higher the $\alpha$, the larger the labor force of the country.</li>
    <li>The higher the $\alpha$, the more labor-intensive a country's production is.</li>
    <li>The higher the $\alpha$, the more capital-intensive a country's production is.</li>
</ol>

Assign the letter corresponding to your answer to `q1_5` below.


```python
q1_5 = ...
```

```python
grader.check("q1_5")
```

### Returns to scale

The significance of the exponents adding up to 1, i.e. $\alpha + \beta = 1$, is that this implies **constant returns to scale**. If all inputs are scaled by a common non-zero factor, the output will be scaled by that same factor. Below is a generalization of this:

$$
\begin{aligned}
Y &= A (c \cdot K)^\alpha (c \cdot L)^{1 - \alpha} \\
&= A c^\alpha K ^ \alpha c^{1 - \alpha}L^{1 - \alpha} \\
&= A c^{\alpha + 1 - \alpha}K^\alpha L^{1 - \alpha} \\
&= c \cdot A K^\alpha L^{1 - \alpha}
\end{aligned}
$$

Thus, any increase in either of the inputs will lead to a 1-1 increase in output. This is a significant assumption to make, as it essentially incentivizes companies to continue to "scale" their production inputs. They are not losing out on how much return is produced - they are getting output that matches exactly what they put into production.

The alternative case is when $\alpha + \beta < 1$. This is called **decreasing returns to scale**, and occurs when a company scales their production inputs by a factor of $c$, but gets a scaling in output that is less than $c$.

The last case is when $\alpha + \beta > 1$. This is called **increasing returns to scale**, and occurs when a company increases their production inputs by $c$, but gets an increase in output that is greater than $c$.

Let us visually examine how values of $\alpha$ and $\beta$ affect output.

```python
alpha_beta_slider = widgets.FloatSlider(min = 0.1, max = 3, step = 0.1, value = 1)
a_b_wid = widgets.interactive(change_alpha_beta, alpha_beta_sum = alpha_beta_slider)
display(a_b_wid)
```

```python
a_b_wid.close()
```

**Question 1.6:** In the hypothetical case where a company can increase inputs to N, a very large amount approaching infinity, all cases of returns to scale result in extremeley high output, approaching infinity. Thus, the owner of the company in such a scenario would not care if their production exhibited decreasing, constant or increasing returns to scale. What would the company prefer?

<ol type="A" style="list-style-type: lower-alpha;">
    <li>As long as the company can increase inputs to N, the owner will recieve the profits from an output approaching infinity anyways.</li>
    <li>Even if the company could increase inputs to N, increasing returns to scale will approach infinite output faster than the other two possibilities and therefore the owner should care about the returns to scale on their production.</li>
    <li>With N inputs, the owner will always have more inputs available to bolster production, regardless of returns to scale.</li>
    <li>It depends on the company - some firms would care about the returns to scale and others wouldn't.</li>
</ol>

Assign the letter corresponding to your answer to `q1_6` below.


```python
q1_6 = ...
```

```python
grader.check("q1_6")
```

---

## Part 2: Real-world Datasets and Jupyter's File System

The second part of this lab will go over how to download a dataset, convert it into a format you can interact with through Jupyter notebboks, and prepare it for analysis. This is an essential skill for future classes and in industry, as you will be able to leverage the flexibility of Python and Jupyter notebooks for any kind of data analysis work.

### Downloading the Dataset

In almost all cases in industry, you will have to find the dataset that can answer the question you are trying to solve. However, as we are introducing you to this process, we will be providing the link to such a dataset.

We will be using "Penn World Table" (PWT), a dataset that catalogues information on relative levels of income, output, inputs and productivity for 182 countries between 1950 and 2019. It was compiled by the University of Groningen.

This is the dataset you will be using for Project 2. This lab will provide a broad overview of the data you will be working with.

Once you have found a dataset, you must be able to download or otherwise convert the file into a format Python can read. The most common and easily-read type is a .csv or Comma Seperated Values file. Data analysis packages used in future classes and industry, like Pandas and datascience, have functions that convert .csv files into Tables. We are going to go through such a process now.

**Question 2.1**
1. Go to: https://www.rug.nl/ggdc/productivity/pwt/ and download the Excel datafile. While Python can read Excel files, it's better able to interpret .csv files as they are simpler.
2. Open the Excel file, select the "Data" sheet, then File > Save As. Type in a file name and change the file format to CSV UTF-8 (Comma-delimited) (.csv). It should be among "common formats". If a warning appears, select "Ok". This is because CSVs can only represent data on a single spreadsheet. As this Excel file has multiple spreadsheets, Excel will only convert the active sheet ("Data" - the one we care about) into a .csv file. 

Alternatively, if you do not have access to Excel, you can upload the datafile to Google Drive and open the Excel file in Google Sheets. You will then select the "Data" tab and File > Download > Comma Separated Values (.csv). Tutorial here if you get stuck: https://imgur.com/a/LFC8e1V

Important: Click on the "Data" tab before downloading as only the active sheet gets converted!

### Uploading a Dataset to Jupyter

If you right-click on the Jupyter logo on the top left of the notebook and select "Open Link in New Tab", this will send you to the Jupyter file structure on DataHub. This is very similar to the file structure on your own computer. Your current location is at the root of this file structure tree. Go ahead and explore - many of the materials from your other classes that use Jupyter notebooks can be seen. Just be careful not to delete or rename files. 

Let's try to navigate to the folder where this notebook is located: `fa25-student > lab > lab04`.

**Question 2.2**

1. Upload the .csv file you just downloaded to the same folder as this notebook. Please name the file `pwt1001.csv`. Uploading is accomplished by dragging and dropping the file into the file explorer window.
2. Assign the variable `file_name` below such that it matches the name of the CSV file that you just uploaded.
3. Run the code cell below.

```python
file_name = ...

### Do not edit the code below ###
try: 
    data = Table.read_table(file_name, sep=";")
    assert len(data) > 50
except:
    data = Table.read_table(file_name, sep=",")
data.show(5)
```

```python
grader.check("q2_2")
```

Congratulations! You've uploaded and read in your first dataset! Below are a few questions that will help you get to know the data you will be working with.

<!-- BEGIN QUESTION -->

**Question 2.3.** Are there missing values in `data`? What transformation could we do on the data to mitigate against this?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.4.** Based on our discussion of the Cobb-Douglas Production Function, which columns would be most beneficial for our analysis?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 2.5.** Create a list of all of the countries in `data` in alphabetical order. Call it `countries`. Think about 2-3 countries you would like to explore in the upcoming project.

_Hint:_ How do we get the unique values in a column?

```python
countries = ...
countries
```

```python
grader.check("q2_5")
```

#### Let's calculate a few summary statistics for one of the variables, `rgdpe` (Expenditure-side real GDP at chained PPPs), by making some simple numpy calls to compare living standards.

**Question 2.6.1.** Calculate the minimum and maximum values of rgdpe and assign them to `minimum_rgdpe` and `maximum_rgdpe` accordingly.


```python
rgdpe_no_nan = ...
minimum_rgdpe = ...
maximum_rgdpe = ...
minimum_rgdpe, maximum_rgdpe
```

```python
grader.check("q2_6_1")
```

**Question 2.6.2.** Calculate the mean, median, 25th percentile, and 75th percentile values of rgdpe and assign them to `mean_rgdpe`, `median_rgdpe`, `twentyfifth_rgdpe`, and `seventyfifth_percentile` accordingly.

Hint: What percentile corresponds to the median value?

```python
mean_rgdpe = ...
median_rgdpe = ...
twentyfifth_rgdpe = ...
seventyfifth_rgdpe = ...

mean_rgdpe, median_rgdpe, twentyfifth_rgdpe, seventyfifth_rgdpe
```

```python
grader.check("q2_6_2")
```

<!-- BEGIN QUESTION -->

**Question 2.7.** Now let's create a histogram to visualize this information! Use the `rgdpe` column from our `data` table to create a histogram with 20 equally spaced bins. 

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.8.** Based on the summary statistics you calculated above, as well as the histogram you just created, do you observe anything of particular interest? Note that there is no 'correct' solution here, we just expect you to reasonably justify your choice.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Let's explore the data further! We will evaluate employment levels in Canada and Mexico as a primer to the project you will be working on.

**Question 2.9** Calculate the mean, median, max, and min values of `emp` for both Canada and Mexico and assign them to the relevant variables below:

```python
canada_emp = ...
mean_canada_emp = ...
median_canada_emp = ...
min_canada_emp = ...
max_canada_emp = ...

mexico_emp = ...
mean_mexico_emp = ...
median_mexico_emp = ...
min_mexico_emp = ...
max_mexico_emp = ...

((mean_canada_emp, median_canada_emp, min_canada_emp, max_canada_emp), (mean_mexico_emp, median_mexico_emp, min_mexico_emp, max_mexico_emp))
```

```python
grader.check("q2_9")
```

<!-- BEGIN QUESTION -->

**Question 2.10.** Based solely on the summary statistics you calculated, do you observe anything of particular interest? Do any of the similarities or differences have anything to do with the metric we are exploring, `emp`? Note that there is no 'correct' solution here, we just expect you to reasonably justify your choice.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### That's all! Get ready for a fun and insightful project!

---

## Credits

We would like to thank Professor Raymond Hawkins for his input and Lecture 4 of his Economics 100B course at UC Berkeley, which informed parts of this notebook.

Further, we reference the University of Groningen's Penn World Table 10.0 Dataset throughout this notebook:

Feenstra, Robert C., Robert Inklaar and Marcel P. Timmer (2015), "The Next Generation of the Penn World Table" American Economic Review, 105(10), 3150-3182, available for download at www.ggdc.net/pwt

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    q1_1 results:
        q1_1 - 1 result:
            ❌ Test case failed
            Trying:
                np.isclose(cobb_douglas(1, 2, 2, 0.5), 2)
            Expecting:
                True
            **********************************************************************
            Line 1, in q1_1 0
            Failed example:
                np.isclose(cobb_douglas(1, 2, 2, 0.5), 2)
            Exception raised:
                Traceback (most recent call last):
                  File "/srv/conda/lib/python3.11/doctest.py", line 1355, in __run
                    exec(compile(example.source, filename, "single",
                  File "<doctest q1_1 0[0]>", line 1, in <module>
                    np.isclose(cobb_douglas(1, 2, 2, 0.5), 2)
                  File "/srv/conda/lib/python3.11/site-packages/numpy/_core/numeric.py", line 2447, in 
```

Tests for q1_1 (points: None)
```
>>> np.isclose(cobb_douglas(1, 2, 2, 0.5), 2)
True
```
```
>>> np.isclose(cobb_douglas(1, 4, 0.5, 0.2, 1.1), 0.6155722066724582)
True
```
```
>>> cobb_douglas(1, 2, 2, 0.5) == cobb_douglas(1, 2, 2, 0.5, 0.5)
True
```


Tests for q1_5 (points: [0, 1])
```
>>> q1_5 in ['a', 'b', 'c', 'd']
True
```
```
>>> q1_5 == 'd'
True
```


Tests for q1_6 (points: [0, 1])
```
>>> q1_6 in ['a', 'b', 'c', 'd']
True
```
```
>>> q1_6 == 'b'
True
```


Tests for q2_2 (points: None)
```
>>> assert type(data) == Table
>>> assert data.num_rows == 12810
>>> assert data.num_columns == 52
```
```
>>> assert data.group('countrycode').num_rows == 183
>>> assert data.group('year').num_rows == 70
```


Tests for q2_5 (points: [0, 1])
```
>>> assert type(countries) == np.ndarray
```
```
>>> assert countries[100] == 'Luxembourg'
>>> assert len(countries) == 183
```


Tests for q2_6_1 (points: [0, 1])
```
>>> assert 0 < minimum_rgdpe and minimum_rgdpe < 40
>>> assert 20000000 < maximum_rgdpe and maximum_rgdpe < 30000000
```
```
>>> assert minimum_rgdpe == 20
>>> assert maximum_rgdpe == 20860506
```


Tests for q2_6_2 (points: [0, 1])
```
>>> assert 200000 < mean_rgdpe and mean_rgdpe < 400000
>>> assert 20000 < median_rgdpe and median_rgdpe < 40000
>>> assert 5000 < twentyfifth_rgdpe and twentyfifth_rgdpe < 7000
>>> assert 150000 < seventyfifth_rgdpe and seventyfifth_rgdpe < 160000
```
```
>>> if isinstance(rgdpe_no_nan, list):
...     rgdpe_no_nan = np.array(rgdpe_no_nan)
... elif isinstance(rgdpe_no_nan, pd.Series):
...     rgdpe_no_nan = rgdpe_no_nan.to_numpy()
... elif isinstance(rgdpe_no_nan, np.ndarray):
...     pass
... else:
...     raise TypeError(f'Unsupported data type: {type(rgdpe_no_nan)}. Expected list, pd.Series, or np.ndarray.')
>>> assert np.isclose(mean_rgdpe, np.mean(rgdpe_no_nan))
>>> assert np.isclose(median_rgdpe, np.percentile(rgdpe_no_nan, 50))
>>> assert np.isclose(twentyfifth_rgdpe, np.percentile(rgdpe_no_nan, 25))
>>> assert np.isclose(seventyfifth_rgdpe, np.percentile(rgdpe_no_nan, 75))
```


Tests for q2_9 (points: [0, 1])
```
>>> assert 0 < mean_canada_emp and mean_canada_emp < 30
>>> assert 0 < median_canada_emp and median_canada_emp < 30
>>> assert 0 < min_canada_emp and min_canada_emp < 30
>>> assert 0 < max_canada_emp and max_canada_emp < 30
>>> assert 0 < mean_mexico_emp and mean_mexico_emp < 30
>>> assert 0 < median_mexico_emp and median_mexico_emp < 30
>>> assert 0 < min_mexico_emp and min_mexico_emp < 30
>>> assert 0 < max_mexico_emp and max_mexico_emp < 60
```
```
>>> assert np.isclose(mean_canada_emp, 11.8857)
>>> assert np.isclose(median_canada_emp, 11.5) or np.isclose(median_canada_emp, percentile(50, canada_emp))
>>> assert np.isclose(min_canada_emp, 6)
>>> assert np.isclose(max_canada_emp, 19)
>>> assert np.isclose(mean_mexico_emp, 25.942857)
>>> assert np.isclose(median_mexico_emp, 23.5) or np.isclose(median_mexico_emp, percentile(50, mexico_emp))
>>> assert np.isclose(min_mexico_emp, 8)
>>> assert np.isclose(max_mexico_emp, 55)
```
